---
title: 'Hello World'
description: 'Python cơ bản'
---

## Hello Python

```yaml
type: NormalExercise 
xp: 100 
key: f2d800ab30   
```


Python là ngôn ngữ thông dịch, được blabla


`@instructions`
- Viết câu lệnh đầu tiên `"Hello  World!"` sử dụng hàm `print(_)`

`@hint`
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa

`@sample_code`
```{python}
print(__)
```
`@solution`
```{python}
print('Hello World!')
```
`@sct`
```{python}
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa
```