---
title: 'Hello World'
description: 'Python cơ bản'
---

## Hello Python

```yaml
type: NormalExercise 
xp: 100 
key: f2d800ab30   
```


Python là ngôn ngữ thông dịch, được blabla


`@instructions`
- Viết câu lệnh đầu tiên `"Hello  World!"` sử dụng hàm `print(_)`

`@hint`
cụm từ `"Hello World!"` phải là một chuỗi

`@pre_exercise_code`

```{python}

```


`@sample_code`
```{python}
print(__)
```
`@solution`
```{python}
print("Hello World!")
```
`@sct`
```{python}
cụm từ `"Hello World!"` phải là một chuỗi
```