---
title: List
description: ""
---

## Khởi tạo list

```yaml
type: NormalExercise
key: 58fdf939e4
xp: 100
```

Khác với kiểu dữ liệu `int`,`bool`, list là 1 kiểu dữ liệu phức hợp, có khả lưu trữ các kiểu dữ liệu khác nhau
``

`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```
