---
title: 'Python list'
description: ""
---

## Tạo list

```yaml
type: NormalExercise
key: 747b234047
xp: 100
```

Khác với `bool`, `int`,`float`,... list là một loại kiểu dữ liệu phức hợp. Ví dụ
```
a = "is"
b = "nice"
my_list = ["my", "list", a, b]
```

`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```
