---
title: 'Load Data'
description: 'Load Data'
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: f2d800ab30   
```





`@instructions`
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa

`@hint`
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa

`@sample_code`
```{python}
print(__)
```
`@solution`
```{python}
print('1234')
```
`@sct`
```{python}
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa
aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa
```